---
evaluation_interval: 1m

rule_files:
  - ../rules/docker.yaml

tests:
  - interval: 1m
    input_series:
      - series: obmondo_monitoring{certname="bantha.enableit", alert_id="monitor::system::docker::status"}
        values: 1x1800
      - series: container_last_seen{certname="bantha.enableit", name="packagesign-script-1"}
        values: '1734500308+35x1800'
      - series: container_last_seen{certname="bantha.enableit", name="jenkins"}
        values: '1734500308+60x1800'

    alert_rule_test:
      - alertname: monitor::system::docker::status
        eval_time: 30m
        exp_alerts:
          - exp_labels:
              severity: critical
              certname: bantha.enableit
              name: packagesign-script-1
              alert_id: monitor::system::docker::status
            exp_annotations:
              summary: 'Docker container packagesign-script-1 is down on host bantha.enableit'
              description: 'Docker container **packagesign-script-1** is down since from 15min on server **bantha.enableit**'
