rule_files:
  - prometheus.yaml

tests:
  # ───────────────────────────────────────────────────────────────────────────────────────────────
  # (1) You install Velero for the first time in your cluster. There'll be no backups.
  #     And the velero_backup_last_successful_timestamp gauge metric will not have any values.
  #     REFERENCE : https://github.com/vmware-tanzu/velero/blob/61adb2a6073229995bceade437d6a3409f7dc8b9/pkg/metrics/metrics.go#L696.
  #
  #     We assume that it takes < 15 minutes for a backup to complete.
  #
  #     Someone must go and trigger the first backup (for each schedule), which must successfully
  #     complete within t = 15m. If it doesn't (as in this case), then we must get alerted.
  #
  #     Once the first backup succeeds, the gauge metric will have a value and the alert will go
  #     away.
  # ───────────────────────────────────────────────────────────────────────────────────────────────
  - interval: 5m
    input_series:
      # NOTE : velero_backup_attempt_total{schedule="daily-schedule"} metric exists.
      #        Which indicates that the daily schedule is enabled.
      - series: 'velero_backup_attempt_total{schedule="daily-schedule"}'
        values: "0+0x3"
    alert_rule_test:
      - eval_time: 16m
        alertname: VeleroNoFirstSuccessfulBackup
        exp_alerts:
          - exp_labels:
              severity: critical
              schedule: daily-schedule
            exp_annotations:
              description: No first successful Velero backup found for daily-schedule.
              summary: No first successful Velero backup found.

  # ───────────────────────────────────────────────────────────────────────────────────────────────
  # (2) The first backup completed successfully at t = 20m. And it's currently t = 25m.
  #     No alert should be triggered for the daily schedule.
  # ───────────────────────────────────────────────────────────────────────────────────────────────
  - interval: 5m
    input_series:
      - series: velero_backup_attempt_total{schedule="daily-schedule"}
        values: "0+0x5"
      - series: velero_backup_last_successful_timestamp{schedule="daily-schedule"}
        values: "NaN+0x3 1200+0x02"
    alert_rule_test:
      - eval_time: 25m
        alertname: VeleroNoFirstSuccessfulBackup
        exp_alerts: []
      - eval_time: 25m
        alertname: VeleroUnsuccessfulBackup
        exp_alerts: []

# ───────────────────────────────────────────────────────────────────────────────────────────────
# (3) The first backup completed successfully at t = 20m. And it's currently t = 1d45m.
